rule_files:
  - ../../monitoring/base/victoriametrics/rules/converted/cilium-alertrule.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'up{job="cilium",instance="10.0.0.1:9090"}'
        values: '0+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: CiliumAgentDown
        exp_alerts:
          - exp_labels:
              job: cilium
              severity: critical
              instance: 10.0.0.1:9090
            exp_annotations:
              runbook: Please consider to find root causes, and solve the problems.
              summary: CiliumAgent has disappeared from Prometheus target discovery.
  - interval: 1m
    input_series:
      - series: 'up{job="cilium",instance="10.0.0.1:9090"}'
        values: '0+0x10'
      - series: 'kube_node_spec_taint{effect="NoExecute", job="kube-state-metrics", key="node.kubernetes.io/unreachable", node="10.0.0.1"}'
        values: '1+0x10'
    alert_rule_test:
      - eval_time: 10m
        alertname: CiliumAgentDown
        exp_alerts: []
  - interval: 1m
    input_series:
      - series: 'up{job="cilium-operator"}'
        values: '0+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: CiliumOperatorDown
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              runbook: TBD
              summary: cilium operator is down.
