rule_files:
  - ../../monitoring/base/victoriametrics/rules/converted/kube-storage-version-migrator-alertrule.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'up{job="kube-storage-version-migrator-migrator"}'
        values: '0+0x60'
    alert_rule_test:
      - eval_time: 60m
        alertname: KubeStorageVersionMigratorMigratorDown
        exp_alerts:
          - exp_labels:
              severity: error
            exp_annotations:
              runbook: TBD
              summary: kube-storage-version-migrator migrator has disappeared.
  - interval: 1m
    input_series:
      - series: 'storage_migrator_core_migrator_migrated_objects{resource="configmaps"}'
        values: '0+0x5 1+0x5'
    alert_rule_test:
      - eval_time: 7m
        alertname: KubeStorageVersionMigratorRunning
        exp_alerts:
          - exp_labels:
              severity: info
              resource: configmaps
            exp_annotations:
              summary: migrating configmaps
