rule_files:
  - ../../monitoring/base/victoriametrics/rules/converted/pvc-autoresizer-alertrule.yaml

tests:
  - name: PVCAutoresizerNotWorking not fired when pvc-autoresizer working.
    interval: 1m
    input_series: 
      - series: 'pvcautoresizer_loop_seconds_total'
        values: '1+1x30'
    alert_rule_test:
      - eval_time: 30m
        alertname: PVCAutoresizerNotWorking
        exp_alerts: []
  - name: PVCAutoresizerNotWorking fired when pvc-autoresizer not working.
    interval: 1m
    input_series: 
      - series: 'pvcautoresizer_loop_seconds_total'
        values: '100+0x30'
    alert_rule_test:
      - eval_time: 30m
        alertname: PVCAutoresizerNotWorking
        exp_alerts:
          - exp_labels:
              severity: critical
              category: storage
            exp_annotations:
              runbook: Please check the log for pvc-autoresizer.
              summary: "pvc-autoresizer not working."
  - name: PVCAutoresizerFailed not fired when pvc-autoresizer succeeded.
    interval: 1m
    input_series: 
      - series: 'pvcautoresizer_failed_resize_total'
        values: '10+0x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: PVCAutoresizerFailed
        exp_alerts: []
  - name: PVCAutoresizerFailed fired when pvc-autoresizer failed.
    interval: 1m
    input_series: 
      - series: 'pvcautoresizer_failed_resize_total'
        values: '0+1x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: PVCAutoresizerFailed
        exp_alerts:
          - exp_labels:
              severity: warning
              category: storage
            exp_annotations:
              runbook: Please check the log for pvc-autoresizer and check connectivity for the api server.
              summary: "pvc-autoresizer failed."
